# CS 3110/5110: Data Privacy
## In-Class Exercise, week of 9/14/2026

```python
# Load the data and libraries
import pandas as pd
import numpy as np
from scipy import stats
#import matplotlib.pyplot as plt

adult = pd.read_csv('https://github.com/jnear/cs3110-data-privacy/raw/main/homework/adult_with_pii.csv')
```

```python
df = adult[['Education', 'Marital Status', 'Target']]
df.head(10)
```

Output:
```
Education         Marital Status Target
0  Bachelors          Never-married  <=50K
1  Bachelors     Married-civ-spouse  <=50K
2    HS-grad               Divorced  <=50K
3       11th     Married-civ-spouse  <=50K
4  Bachelors     Married-civ-spouse  <=50K
5    Masters     Married-civ-spouse  <=50K
6        9th  Married-spouse-absent  <=50K
7    HS-grad     Married-civ-spouse   >50K
8    Masters          Never-married   >50K
9  Bachelors     Married-civ-spouse   >50K
```

## Question 6

Imagine the column `Target` is *not* a quasi-identifier, and we generalize the dataset to achieve $k$-Anonymity for $k=2$ as follows:

- Replace each education level below "HS Grad" with `< HS` and others with `>= HS`
- Replace marital status with `Married` and `Not Married`
- Delete rows as required to achieve $k$-anonymity for $k=2$

For which rows is a homogeneity attack possible, and why?

<!-- grade_id: cell-5932f02f09b5cacb, points: 5, grade: true, solution: true -->
If any group contains only a single target than a homogenaity attack is possible.
For k = 2, the combo < HS and married only has one target value

## Question 7

Write code to generalize the `Zip` column in the adult dataset.

<!-- grade_id: cell-be9a4264a0e58938, points: 0, grade: false, solution: true -->
```python
# replace l least significant digits with 0s
#05482 with l = 2 -> 05400
def generalize_zip(zip, l):
    new_zip = int(zip/ 10**l) * 10**l
    return new_zip
```

<!-- grade_id: cell-256688fe78b8f7bf, points: 5, grade: true, solution: false -->
```python
# Test cases
assert generalize_zip(47401, 0) == 47401
assert generalize_zip(47401, 1) == 47400
assert generalize_zip(47401, 2) == 47400
assert generalize_zip(47401, 3) == 47000
assert generalize_zip(47401, 4) == 40000
assert generalize_zip(47401, 5) == 0
```

## Question 1

Write a *counting query* to determine whether or not Karrie Trusslove's age is 39.

<!-- grade_id: cell-e9c0c06eb7735ea3, points: 0, grade: false, solution: true -->
```python
def karrie_query():
    return len(adult[(adult['Name'] == 'Karrie Trusslove') & (adult['Age'] == 39)])
```

<!-- grade_id: cell-f93cfd0eb59130ed, points: 5, grade: true, solution: false -->
```python
# TEST CASE 

assert karrie_query() == 1
```

## Question 2

Add Laplace noise to the counting query you wrote in the last question to ensure differential privacy for $\epsilon = 1.0$.

<!-- grade_id: cell-8ad72c55497bc534, points: 0, grade: false, solution: true -->
```python
# goal add random noise to the 0 or 1 I get from the karrie query
# to obscure whether the answer is 0 or 1
# this prevents the adversary from knowing Karrie's age
# Problem: this means we return the "wrong answer"
def dp_karrie_query():
    original_result = karrie_query()
    epsilon = 0.01
    noise = np.random.laplace(loc=0, scale = 1/epsilon)
    #noise = np.random.randint(1,10)
    return original_result + noise
```

```python
dp_karrie_query()
```

Output:
```
337.3127884821835
```

<!-- grade_id: cell-a7c9bd084e5c5bdf, points: 5, grade: true, solution: false -->
```python
# TEST CASE 
epsilon = 0.1
q2_runs = [dp_karrie_query() for _ in range(100)]
noise_runs = [np.random.laplace(loc=1, scale=1/epsilon) for _ in range(100)]

assert stats.wasserstein_distance(q2_runs, noise_runs) < 1
```

Output:
```
AttributeError                            Traceback (most recent call last)
Cell In[58], line 6
      3 q2_runs = [dp_karrie_query() for _ in range(100)]
      4 noise_runs = [np.random.laplace(loc=1, scale=1/epsilon) for _ in range(100)]
----> 6 assert stats.wasserstein_distance(q2_runs, noise_runs) < 1

AttributeError: module 'scipy.stats' has no attribute 'wasserstein_distance'
```

## Question 3

In 2-4 sentences, describe how `dp_karrie_query` protects Karrie's privacy.

<!-- grade_id: cell-07337a085c0ba0e1, points: 5, grade: true, solution: true -->
- The original query is adversarially constructed to violate Karrie's privacy
- The answer is 1 if Karrie is 39 and 0 otherwise
- The noise introduces uncertainty for an adversary who wants to guess whether the answer was 0 or 1 based on the DP result
- More noise (lower epsilon, higher scale parameter) makes this uncertainty greater
- The noise4 (and thuse the uncertainty) is independent of the data (and anything else in the world), so the adversary can't undo the noise using aux. informaiton. 

## Question 4

Implement a function to construct an *empirical distribution* over random samples from the Laplace distribution with location 0 and scale `scale`. Use 1000 samples.

<!-- grade_id: cell-e7ddf0e4b195dfdc, points: 0, grade: false, solution: true -->
```python
def empirical_dist_lap(scale):
    samples = [np.random.laplace(loc=0, scale=scale) for _ in range(1000)]
    
    raise NotImplementedError()

epsilon1 = 1.0
epsilon2 = 2.0

plt.hist(empirical_dist_lap(1/epsilon1), bins=50);
plt.hist(empirical_dist_lap(1/epsilon2), bins=50, alpha=.7);
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[61], line 9
      6 epsilon1 = 1.0
      7 epsilon2 = 2.0
----> 9 plt.hist(empirical_dist_lap(1/epsilon1), bins=50);
     10 plt.hist(empirical_dist_lap(1/epsilon2), bins=50, alpha=.7);

NameError: name 'plt' is not defined
```

<!-- grade_id: cell-e68891bb70939fd0, points: 5, grade: true, solution: false -->
```python
# PLACEHOLDER for test cases

```

## Question 5

Consider the program below.

```python
counting_query = lambda: 1

result_1 = counting_query() + np.random.laplace(loc=0, scale=1/epsilon1)
result_2 = counting_query() + np.random.laplace(loc=0, scale=1/epsilon2)
result_3 = counting_query() + np.random.laplace(loc=0, scale=1/epsilon1)

print(result_1, result_2, result_3)
```

Write a function `total_epsilon` that returns the total privacy cost for this program.

<!-- grade_id: cell-32b1cdb0a2990e8f, points: 0, grade: false, solution: true -->
```python
def total_epsilon():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e484920fc1429ad0, points: 5, grade: true, solution: false -->
```python
# Placeholder for test case

```